---
title: "Statistical testing"
author: "Simon Thornley"
date: "`r format(Sys.time(), '%e %B, %Y')`"
output: 
  learnr::tutorial:
        progressive: true
        allow_skip: true
  prettydoc::html_pretty:
                 theme: tactile
runtime: shiny_prerendered
description: >
  Inspecting our data for errors!
editor_options: 
  markdown: 
    wrap: 72
---

```{r setup, include=FALSE}
library(learnr)
library(tufte)
library(rio)
library(visdat)
library(fs)
library(ggplot2)
library(DiagrammeR)
library(r2symbols) # cheatsheet here: https://r2symbols.obi.obianom.com/symbols/
library(learnr)
library(gradethis)
library(skimr)
library(lattice)
library(epiDisplay)
library(epitools)
library(arsenal)
gradethis::gradethis_setup()
load("./data/df.rdata")
df <- df_no_dups
df$Case_status_f <- factor(df$Case_status, levels = c("Control", "Case"))
df[df$Mother_smoke == "0", "Mother_smoke"] <- "No"

#df <- rio::import("https://flexiblelearning.auckland.ac.nz/data-analysis/menu/1/files/simple_sids_epiinfo2.xlsx")
#dir.create("./data")
#rio::export(df, "./data/sids.xlsx")

# tutorial_options(
#   exercise.timelimit = 60,
#   # A simple checker function that just returns the message in the check chunk
#   exercise.checker = function(check_code, ...) {
#     list(
#       message = eval(parse(text = check_code, ...)),
#       correct = logical(0),
#       type = "info",
#       location = "append"
#     )
#   }
# )
knitr::opts_chunk$set(error = TRUE)
```

<br> </br>

<center>![The aim is to keep infants alive!](images/colin-maynard-CEEhmAGpYzE-unsplash.jpg "Photo by Colin Maynard on Unsplash"){#id
.class width="40%" height="40%"}</center>

# Introduction

In the [last
session](https://sithor.shinyapps.io/Tutorial_3/){target="_blank"}, we
discussed how to import data into
[`R`](https://cran.r-project.org/){target="_blank"} and do some
preliminary checks prior to analysis.

Now we are ready for analysis.

We will learn how to:

-   visualise the relationship between a categorical and continuous variable using boxplots and histograms
-   conduct *t*-tests of significance for these relationships
-   Interpret the results of such tests.

This session, we will focus on continuing to take baby steps with `R`
and revise and develop some of the concepts from the last tutorial.

As a little bit of revision for last week, here's a quick multi-choice
question about checks before analysis of epidemiological data.

## A little revision

* Attempt the following question...

```{r revision, echo=FALSE}
question("What are the three most important issues to consider when you import data into R?",
  answer("Statistical significance, exposure and confounding variables are present.", message = "These are useful to know, but not crucial at the first stage!"),
  answer("Format of data, number of observations, type of study", message = "These are important, but not crucial for avoiding errors!"),
  answer("Length of dataset, number of variables, missing values", message = "One is correct... the first two are useful to know, but not crucial"),
  answer("Duplicates, missing values and range checks", correct = TRUE, message = "Well done, these are important to consider"),
  allow_retry = TRUE
)
```

## The cot-death case-control study

Cot death was an important problem in the 1980s and New Zealand
experienced an epidemic of the disease. There was a competition between
the lab scientists and epidemiologists over who would be able to come up
with an answer to respond to the epidemic.

The critical elements of a case-control study are the definition of
cases and controls (the disease), and the potential exposures.

<center>![Case-control
study](images/download.png "Case-control")</center>

Here, **cases** were mothers of babies who had a child that died from
unexplained causes in the first year of life, whereas **controls** were
mothers of children less than one year old who were in normal health.

Some of the possible **causes** or **exposures** were:

-   sleeping position of the baby on the night before death, or randomly
    chosen for controls (front vs. back or side)
-   bed-sharing with parents (yes or no)
-   mother smoked cigarettes (yes or no)

```{r case-control, echo = FALSE}
question("Why do you think a *case-control study* was conducted to investigate risk factors for cot-death?",
  answer("Cot death is a rare disease", correct = TRUE, message = "Yes, case-control studies are efficient designs for studying rare diseases."),
  answer("Maternal smoking was rare", message = "This was not the case, but cohort studies are, in any case, better for studying rare diseases."),
  answer("The design adjusts for confounders", message = "No, only randomisation will do this! That is, a randomised controlled trial is needed to accomplish this."),
  allow_retry = TRUE
)
```

## Loading libraries

The main library that will help us achieve our goals today will be:

```{r eval = FALSE}
lattice  # plots!
epitools # calculating odds ratios & risk ratios!
arsenal # making nice tables!
```

which is useful for making some plots.

To install and load libraries in R, please complete the following
exercise, to install and load the `lattice` package.

**Note**: all the other libraries are loaded, but in the real world, if we were using all three libraries, we would need to load all three in the same manner.

```{r load_epiDisplay, exercise = TRUE}
if (!require()) install.packages()
library()
```

```{r load_epiDisplay-solution}
if (!require(lattice)) install.packages("lattice")
library(lattice)
```

<div id="load_epiDisplay-hint">
**Hint:** You need to include quotes (`""`) around and `lattice`, for
`install.packages()`, but not for `require` or `library` - they don't
need quotes.
</div>

```{r load_epiDisplay-check}
grade_this_code("Well done, it is a vital skill to install and load new libraries in `R`!")
```

The data is already loaded into this webpage, and checked its quality,
so since we've covered that last session, we won't do it now, but here's
a [link](https://sithor.shinyapps.io/Tutorial_3/){target="_blank"} just
in case you've forgotten. Remember, we called the spreadsheet (in R
`data.frame`) "`df`" short for "`data.frame`".

## Does birth weight influence cot-death?

We want to know whether (or not) birth weight is likely to influence risk of
cot-death.

If birth weight has an influence, then we expect to see a difference in
the *mean* birth weight between the cases and controls.

If there is no influence, we expect the means in the two groups
to be very similar.

How do we know how different the means must be to not just be due to
chance?

## First step: visualise

To do this we compare the means of the birth weights between the cases
and the controls. Let's first examine the distribution of birth weight
using the trusty histogram.

### Histograms

Note: here, we are using the `histogram()` function from the `lattice`
package.

The `lattice::histogram()` function tells R to use the `histogram()`
function from the `lattice` package.

Once you've run the code below, change it to plot a histogram of
`Mother_age`. Another continuous variable.

Notice that the first argument is a
[`formula`](http://www.mysmu.edu/faculty/anthonytay/MFE/01_R_Formulas.pdf){target="_blank"}
which uses the tilde (`~`) to separate variables, usually *outcomes* on
the left and *exposures* on the right. It is often used for regression,
which we will cover later. It allows the use of functions in a compact
and symbolic form.

We can find out more about how to use the formula in this function by
using the following code:

```{r plot_cont_help, exercise = TRUE}

?lattice::histogram

```

Scroll down to the first argument, `x` which is a `formula` of the form
`~ x | g1` where:

-   `x` is the numeric variable, and
-   `g1` is the factor variable.

Notice, that the help file specifies that `x` should be a number and
`g1` should be a `factor`.

Here, `g1` will be `Case` status.

Also notice that the second argument specifies the data.frame we are
using (`df`).

Run the given code and then change it to plot a histogram of
`Mother_age`. Another continuous variable in the data.

```{r plot_cont, exercise = TRUE}

lattice::histogram(~ Birth_wt, data = df)

```

```{r plot_cont-solution}
lattice::histogram(~ Mother_age, data = df)
```

<div id="plot_cont-hint">
**Hint:** You need to replace `Birth_wt` for the other variable
`Mother_age`.
</div>

```{r plot_cont-check}
grade_this_code("Well done, it is a vital skill to make histograms in R")
```

And then plot the continuous variable (`Birth_wt`) by `Case_status`.

-   First run the example, then
-   change the code below to plot `Mother_age` by `Case_status`.

Note that `xlab` is short for x-axis label.

```{r by_group, exercise = TRUE}

lattice::histogram(~ Birth_wt | Case_status, 
                   data = df[, c("Birth_wt", "Case_status")] |> na.omit(),
                   xlab = "Birth weight (grams)",
                   layout = c(1, 2),
                    panel = function(x, ...) {
                    panel.histogram(x, ...)
                    panel.abline(v = mean(x, na.omit = TRUE), 
                                 col = 'black', lty = 2 , lwd = 3)
                    } )

```

```{r by_group-solution}

lattice::histogram(~ Mother_age | Case_status, 
                   data = df[, c("Mother_age", "Case_status")] |> na.omit(),
                   xlab = "Mother age (years)",
                   layout = c(1, 2),
                    panel = function(x, ...) {
                    panel.histogram(x, ...)
                    panel.abline(v = mean(x, na.omit = TRUE),
                                 col = 'black', lty = 2 , lwd = 3)
                    } )
```

<div id="by_group-hint">
**Hint:** You need to replace `Birth_wt` for the other variable
`Mother_age` where it is present in the code. Also, change the label for
the x-axis (`xlab` to something more appropriate `Mother age (years)`).
</div>

```{r by_group-check}
grade_this_code("Well done, it is a vital skill to make histograms by a categorical variable in R")
```

#### Test your knowledge

```{r int_hist, echo = FALSE}
question("Which group has the largest birth weight babies on average?",
  answer("Controls", correct = TRUE, message = "Yes, the dashed black line (group mean) is further to the right for controls compared to cases."),
  answer("Cases", message = "No, the dashed black line (group mean) for cases is to the left of that for controls."),
  answer("No difference", message = "We can only make this conclusion if the statistical test is not significant. We will do this next!"),
  allow_retry = TRUE
)
```

```{r int_hist_2, echo = FALSE}
question("Which group has the oldest mothers on average?",
  answer("Controls", correct = TRUE, message = "Yes, the dashed black line (group mean) is further to the right for controls compared to cases."),
  answer("Cases", message = "No, the dashed black line (group mean) for cases is to the left or lower than for controls."),
  answer("No difference", message = "We can only make this conclusion if the statistical test is not significant.We will do this next! Try again."),
  allow_retry = TRUE
)
```

How large is the difference in birth weight and maternal age by
case-status?

<style>
div.blue { background-color:#e6f0ff; border-radius: 5px; padding: 20px;}
</style>
<div class = "blue">

Note: there is some complex code here to explain.

-   The argument, `layout = c(1, 2)` changes the orientation of the two
    plots from side by side (`layout = c(2, 1)`; one row two columns) to
    one on top of the other (one column, two rows).

-   This histogram function throws a wobbly (error) if there is missing
    data. The `df[, c("Birth_wt", "Case_status")] |> na.omit()` subsets
    the data to only the data we are using and removes missing data.
    This is a fix for many **annoying** issues in `R`!

-   The rest of the code:<br></br> `layout = c(1, 2),` <br></br>
    `panel = function(x, ...){` <br></br> `panel.histogram(x, ...)`
    <br></br> `panel.abline(v = mean(x, na.omit = TRUE),` <br></br>
    `col = 'black', lty = 2 , lwd = 3)` <br></br> `}`<br></br> plots the
    mean so that we can compare the magnitude of the difference in
    means, which is what our *t*-test will assess. I cheated and grabbed
    it from [stack overflow](https://stackoverflow.com/){target="_blanck"} and adapted. 
    Don't worry about it too much.
    It just works!! The idea here is to adapt the code we have to other
    circumstances. Often just a small tweak is needed!

-   Notice how each argument is on a separate line. This helps the code
    not run off the page, so it makes it easier to read!

</div>


### The boxplot

Whilst the histogram is a useful tool, you can see it is a bit of a pain
to compare measures of central tendency, such as means and medians. The
*boxplot* is very useful for comparing medians (central tendency) and
the distribution of continuous variables, without too much complexity.

We will redo what we've done comparing birth weight by disease (SIDS)
status. Run this code and change and submit the code to check maternal
age `Mother_age` by case status.

```{r box, exercise = TRUE}
boxplot(Birth_wt ~ Case_status, 
        data = df)

```

```{r box-solution}
boxplot(Mother_age ~ Case_status, 
        data = df)
```

<div id="box-hint">
**Hint:** You need to replace `Birth_wt` for the other variable
`Mother_age`.
</div>

```{r box-check}
grade_this_code("Well done, it is a vital skill to be able to make a boxplot in R.")
```

#### Interpret the boxplot

A boxplot has a number of features that are important to recognise.

<center>![Source:
https://www.kdnuggets.com/2019/11/understanding-boxplots.html](images/1_NRlqiZGQdsIyAu0KzP7LaQ.png "boxplot and corresponding normal distribution"){.class
width="40%" height="40%"}</center>

The boxplot divides the continuous variable into four quartiles, with
the second and third quartiles represented by the box, divided by the
median or middle value. The whiskers are the last observation within 1.5
x the interquartile range (length of box). Any values beyond 1.5 x the
interquartile range are plotted separately as possible outliers.


Enter the median value of the maternal age in years of **cases**.

```{r int_box, exercise = TRUE}

```

```{r int_box-solution}
25
```

<div id="int_box-hint">
**Hint:** You need to check the value on the vertical (y) axis of the middle line in the box above `Case`.
</div>

```{r int_box-check}
grade_this( {
if(isTRUE(all.equal(.solution, .result, tolerance = 0.1) )){
    pass("Great work! You are getting good at interpreting box plots")
  }
  fail("Please revise the meaning of the boxplot in the previous section and have another go!")
})
```

Enter the median value of the maternal age in years of **controls**.

```{r intc_box, exercise = TRUE}

```

```{r intc_box-solution}
27.5
```

<div id="intc_box-hint">
**Hint:** You need to check the value on the vertical (y) axis of the middle line in the box above `Controls`.
</div>

```{r intc_box-check}
grade_this( {
if(isTRUE(all.equal(.solution, .result, tolerance = 0.2) )){
    pass("Great work! You are getting good at interpreting box plots")
  }
  fail("Please revise the features of a boxplot in the previous section and have another go!")
})
```


Enter the median difference of maternal age in years between **cases** and **controls**.

```{r intd_box, exercise = TRUE}

```

```{r intd_box-solution}
2.5
```

<div id="intd_box-hint">
**Hint:** You need to check the difference in y-value between the middle lines for `cases` and `controls`.
</div>

```{r intd_box-check}
grade_this( {
if(isTRUE(all.equal(.solution, .result, tolerance = 0.2) )){
    pass("Great work! You are getting very good at interpreting box plots!")
  }
  fail("Please revise the meaning of the boxplot in the previous section and have another go!")
})
```


## The *t*-test: is the difference in group means 'significant'? 

<center>![Due to chance or
not?](images/alperen-yazgi-QuP5RL_E5oE-unsplash.jpg "Photo by Alperen Yazgi on Unsplash"){.class
width="40%" height="40%"}</center>

We now see that there is a difference between birth weight and maternal age between cases and controls, but is this likely to be due to chance or not? The plot does not tell us this. The plot checks the symmetry of the distribution and the magnitude of the difference. The statistical test tells us whether this difference is due to chance or not. Is the difference likely to be **true** or is it likely to be **due to chance**?

To check whether or not this difference is due to chance (the *null
hypothesis*), we need to do a *t*-test.

The `t.test(y ~ x)` function in R is used where `y` is numeric and `x`
is a two-level category.

It will estimate the means in both groups, the magnitude of the
difference in means, and the $P$-value which assesses whether the
difference is likely to be due to chance or not.

-   \< 0.05 = difference in means is unlikely to be due to chance or
    *statistically significant* (consider type-1 error also) = evidence
    for association and possible causal influence.
-   `r sym("gtequal")` 0.05 = difference in means is likely to be due to
    chance or *not statistically significant* (consider type-2 error) =
    evidence for no association and no likely causal influence.

Run the following code, and adapt it to test for a difference in
`Mother_age` by `Case_status`.

Note the use of the `formula` argument.

```{r t_test, exercise = TRUE}
t.test(Birth_wt ~ Case_status, 
       data = df)

```

```{r t_test-solution}
t.test(Mother_age ~ Case_status, 
       data = df)
```

<div id="t_test-hint">
**Hint:** You need to replace `Birth_wt` with the other variable
`Mother_age`.
</div>

```{r t_test-check}
grade_this_code("Well done, it is a vital skill to be able to do *t*-tests in R.")
```

**Note**: `2.2e-16` is **scientific notation** which is the same as
$2.2 \text{ x } 10^{-16}$ or 0.00000000000000022. This is a very small
number, which means that if there were truly no difference in birth
weight between the two groups, it would be extremely rare to see such a
difference that we have observed, if the study were repeated over and over. 
In fact, it would be expected to get such a result, or more extreme, less than one in **one
hundred trillion** repetitions!


As a bit of revision, we will remind ourselves of what the *t*-test
means.

The *t* statistic estimates how many standard errors the observed
results are from the null value of no difference.

The formula for the *t* statistic is: \begin{align}
        t  &= \frac{\text{difference in means}}{\text{standard error of difference in means}} \\
           &= \frac{\bar{x}_1 - \bar{x}_2}{\text{standard error of difference in means}} \\
           &= \frac{\bar{x}_1 - \bar{x}_2}{\sqrt{\frac{\sigma_1^2}{n_1} + \frac{\sigma_2^2}{n_2}}} \\
\text{Where:}\\
\sigma &= \text{the standard deviation of the sample, in groups 1 and 2} \\
n &= \text{the total number of measurements taken in groups 1 and 2} \\
\bar{x} &= \text{the mean of groups 1 and 2} \\
\end{align}

The critical threshold for the *t* statistic varies according to the
degrees of freedom which is $n_1$ + $n_2$ - 2.


We can visualise this distribution with the following code:

```{r eval = TRUE}
x <- seq(-10, 10, length = 100)
dfreedom = 491
plot(x, dnorm(x), type = "l", lty = 2, xlab = "t-value (standard errors from the mean)", 
     ylab = "Density", 
     main = "Comparison of t-distributions", col = "blue", lwd = 1.5)
    lines(x, dt(x, dfreedom), col  = "red")
    abline(v = c(-9.4866, 9.4866), col = "green", lwd = 3)
```

The [green]{style="color: green;"} line shows how likely the observed
result is under the **null** (or no association) hypothesis. You can see that once we are beyond about 2.5 standard errors from the null, we are well outside the bounds of what is likely to happen from chance. At 9.5 standard errors, we are deeply in very improbable territory for the null hypothesis belief.

This tells us that there is likely to be a real or true difference and thus an **association** between **birth weight** and **maternal age** and **SIDS status**. This may be a direct cause or may be explained by other *confounding* factors, or other *biases*, as we shall discuss in future lectures and tutorials. 

## Categorical tests of significance
<center>![Source:
https://unsplash.com/s/photos/baby-sleeping](images/igordoon-primus-Pb0NiOMae0w-unsplash.jpg "baby sleeping"){.class
width="40%" height="40%"}</center>
Let's say we are interested in whether or not sleeping an infant on their front is a risk-factor for cot-death (`Case_status`). Our relevant exposure variable here is `Sleep_position`. We can look at the association between `Case_status` and `Sleep_position`, with the following code:

```{r cc_test, exercise = TRUE}
## Estimate risk ratios and odds ratios
tab <- table(df$Sleep_position, df$Case_status)
epitools::oddsratio(tab)
epitools::riskratio(tab)
```


```{r ini_cat_assoc, echo = FALSE}
question("For a case-control study, which measure of association is most appropriate?",
  answer("Risk ratio",  message = "Sorry, complex arguments related to conditional probability mean that risk ratios are not appropriate for case-control studies."),
  answer("Odds ratio", correct = TRUE, message = "Yep, this is the one!"),
  allow_retry = TRUE
)
```



### What does the measure of association mean?

Hold on just a minute! Check your output and interpret the answer. It is important to be able to do a manual calculation to check on what the computer is doing! 

Computers are great at doing calculations, but they can inadvertently trip us up if we are not careful!

```{r cat_ass, echo = FALSE}
question("What is your interpretation of the output?",
  answer("Infants who sleep face down are at greater risk of cot-death than children who are slept on their back.", correct = TRUE, message = "Yes, this is true, although at first glance, it doesn't appear so. Well done!"),
  answer("Infants who sleep face down are at lower risk of cot-death than children who are slept on their back", message = "No, you need to carefully look at the numbers. We will estimate the odds ratio now."),
  answer("There is no difference in risk of cot-death between children slept on their front and back.", message = "No, the *P*-value is less than 0.05 meaning that the association is statistically significant, so this is not the correct conclusion."),
  allow_retry = TRUE
)
```

```{r in_cat_assoc, echo = FALSE}
question("What is the **odds ratio**, comparing the odds of cot-death in infants sleeping face down to on their back?",
  answer("(13/(32 + 13)) / (18/(18 + 220)) = 3.8", message = "Nearly there. This is a **risk ratio** rather than an **odds ratio**."),
   answer("(13/32)/(18/220) = 5.0", correct = TRUE, message = "Yes, this is true, although at first glance, it doesn't appear so. Well done!"),
  answer("(32/13)/(220/18) = 0.20", message = "No, this is the odds ratio of cot-death, comparing back  with face-down sleepers."), 
  answer("(32/(32 + 13)) / (220/(18 + 220)) = 0.77", message = "No, this is the risk ratio of cot-death, comparing back with face-down sleepers."),
  allow_retry = TRUE
)
```

To get the computer to give us the correct output, we will need to re-order the outcome variable. This is accomplished by turning `Case_status` into a *factor* variable.

The following code may be used:
```{r initi_cat_assoc, exercise = TRUE}
## reorder Case_status variable!
df$Case_status_f <- factor(df$Case_status, levels = c("Control", "Case"))

## Make a 2 x 2 table
tab <- table(df$Sleep_position, df$Case_status_f)

## Estimate odds ratio
epitools::oddsratio(tab)

## Estimate risk ratio
epitools::riskratio(tab)
```

Have a go at now estimating the association between `Case_status_f` and `Mother_smoke`.

```{r cc, exercise = TRUE}
tab <- table()
epitools::oddsratio()
```

```{r cc-solution}
tab <- table(df$Mother_smoke, df$Case_status_f)
epitools::oddsratio(tab)
```

<div id="cc-hint">
**Hint:** You need to use `Case_status_f` as the **exposure** and `Mother_smoke` as the outcome. 
</div>

```{r cc-check}
grade_this_code("Brilliant! Well done!")
```

## Tabulating results in publications

Often, we like to see the association between a number of sociodemographic factors and case and control status in a table.

```{r table1, exercise = TRUE}
## Use attributes to make nice labels

attr(df$Mother_age,'label')  <- 'Maternal age (years)'
attr(df$Mother_smoke,'label')  <- 'Maternal smoking status'

## Automagically make a table 1
tab <- arsenal::tableby(Case_status ~ Mother_age + Mother_smoke + Occupation,
                        data = df, total = TRUE, ordered.simplify = TRUE)

summary(tab, text = TRUE)

```


```{r int_tab, echo = FALSE}
question("Find the best interpretation of the table above.",
  answer("All three variables are associated with case-status and older mothers are particularly at risk of cot-death", message = "Sorry, it is actually younger mothers that at risk."),
   answer("No variables are associated with cot-death", message = "Ooops. Check the *P*-values. They are all associated with cot-death!"),
  answer("All three variables are associated with case-status and younger mothers, smokers and 'low occupation' mothers are particularly at risk of cot-death", correct = TRUE, message = "Bang on!"), 
  allow_retry = TRUE
)
```


Have a go now at adding `Sleep_position` and `Bedshare` to the table you created earlier.
Don't worry about changing the attributes.
```{r tab1, exercise = TRUE}
## Use attributes to make nice labels

attr(df$Mother_age,'label')  <- 'Maternal age (years)'
attr(df$Mother_smoke,'label')  <- 'Maternal smoking status'

## Automagically make a table 1, add `Sleep_position` and `Bedshare` 
tab <- arsenal::tableby(Case_status ~ Mother_age + Mother_smoke + Occupation,
                        data = df, total = TRUE, ordered.simplify = TRUE)

summary(tab, text = TRUE)

```


```{r tab1-solution}
attr(df$Mother_age,'label')  <- 'Maternal age (years)'
attr(df$Mother_smoke,'label')  <- 'Maternal smoking status'
## Automagically make a table 1, add `Sleep_position` and `Bedshare`
tab <- arsenal::tableby(Case_status ~ Mother_age + Mother_smoke + Occupation +
                          Sleep_position + Bedshare,
                        data = df, total = TRUE, ordered.simplify = TRUE)

summary(tab, text = TRUE)
```

<div id="tab1-hint">
**Hint:** You might want to add ` + Sleep_position + Bedshare` to the `tableby()` function call. 
</div>


```{r tab1-check}
grade_this_code("Superb! Well done!")
```

### Homework

Try and emulate these analyses in a desktop version of `posit`. 

- Import the data
- Make both histograms and boxplots which illustrate whether there is a difference in `Gestation` (time from conception to birth) between cases and controls? 
- Test for statistical significance between cases and controls for `Gestation`.
- Could such an association explain the observed differences in birth weight?
- Make an appropriate table of these results for publication!
